<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 573031  -->
  <question type="stack">
    <name>
      <text>P05 - Dígitos con restricciones (Par/Impar)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>¿Cuántos números enteros de {@N@} cifras se pueden formar, si el primer dígito (de izquierda a derecha) debe ser impar y el último dígito debe ser par?</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
      <text>/* Generamos N entre 3 y 7 para que el número no sea inmensamente grande pero sí interesante */
N: 3 + rand(5);

/* Cálculo de la respuesta:
   Primer dígito (impar): 5 opciones
   Último dígito (par): 5 opciones
   Dígitos intermedios (N-2): 10 opciones cada uno */

ta1: 5 * 5 * (10^(N-2));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>N={@N@}, resp={@ta1@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><i class="fa fa-check"></i> Respuesta correcta, bien hecho.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><i class="fa fa-adjust"></i> Su respuesta es parcialmente correcta.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><i class="fa fa-times"></i> Respuesta incorrecta.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>Escribe tu respuesta aquí</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords>ta1</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>¡Correcto! Muy bien.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p data-path-to-node="15,1,0">Incorrecto. Vamos por partes:</p>
<ol start="1" data-path-to-node="15,1,1">
<li>
<p data-path-to-node="15,1,1,0,0">Para el primer dígito (impar) tienes 5 opciones: {1, 3, 5, 7, 9}.</p>
</li>
<li>
<p data-path-to-node="15,1,1,1,0">Para el último dígito (par) tienes 5 opciones: {0, 2, 4, 6, 8}.</p>
</li>
<li>
<p data-path-to-node="15,1,1,2,0">Te quedan {@N-2@} dígitos en medio, donde cada uno tiene 10 opciones (0-9).</p>
</li>
</ol>
<p data-path-to-node="15,1,2">El cálculo es: 5 * 10^{@N-2@} * 5.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>